# Denison CS181/DA210 Homework 1.f - Step 2

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path

datadir = "publicdata"
```

---

## Part B

**Q4:** Assume that `path` refers to a CSV file with the same format used in our variations of `topnames` with a header line containing the four column names, and then rows of comma-separated field values.  Write a function

```
    readTopNamesLoD(path)
```

that reads the file and creates a LoD representation and returns the structure.

<!-- grade_id: cell-3ac611d5fbae072d, points: 0, grade: false, solution: true -->
```python
def readTopNamesLoD(path):
    """ 
    This function reads the CSV file and creates a LoD 
    representation and returns the structure
    """
    new_list = []
    with open (path, 'r') as filedata:
        header = filedata.readline().strip().split(',')
        line = filedata.readline()
        while line != "":
            data_line = line.strip().split(',')
            data = {header[0]: int(data_line[0]), header[1]: data_line[1], header[2]: data_line[2], header[3]: int(data_line[3])}
            new_list.append(data)
            line = filedata.readline()
    return new_list

```

```python
# Debugging cell
tn10data = readTopNamesLoD(os.path.join(datadir, "tn10.csv"))
print(tn10data)
```

Output:
```
[{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]
```

<!-- grade_id: cell-ab5935ca57c293ca, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert readTopNamesLoD.__doc__ is not None # don't forget the docstring!

tn10data = readTopNamesLoD(os.path.join(datadir, "tn10.csv"))

assert isinstance(tn10data, list)
assert len(tn10data) == 10

assert len(tn10data[0]) == 4
assert isinstance(tn10data[0], dict)

assert tn10data[0]['sex'] == 'Female'
assert tn10data[0]['year'] == 2014
assert tn10data[0]['name'] == 'Emma'
assert tn10data[0]['count'] == 20936
```

**Q5:** Write a function 

```
    filterTopNamesLoD(tnLoD, threshold)
```

to create a **filtered** copy of a topnames LoD `tnLoD` (with columns `year`, `sex`, `name`, `count`) so that only rows with a `count` value greater than or equal to `threshold` are present in the newly created LoD.

Your function should return the new and filtered LoD, nd not modify the original `tnLoD`.

```python
def filterTopNamesLoD(tnLoD, threshold):
    """ 
    This function create a filtered copy of a topnames 
    LoD so that only rows with count>=threshold are 
    present in a new LoD
    """
    return [new_dict for new_dict in tnLoD if new_dict['count'] >= threshold]
```

```python
# Debugging cell
tn10 = readTopNamesLoD(os.path.join(datadir, "tn10.csv"))
tn10_filter = filterTopNamesLoD(tn10, 19000)
print(tn10_filter)
```

Output:
```
[{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]
```

```python
# Testing cell
assert filterTopNamesLoD.__doc__ is not None # don't forget the docstring!

LoD = [{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, 
       {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, 
       {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, 
       {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, 
       {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, 
       {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, 
       {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, 
       {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, 
       {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, 
       {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]

filtered = filterTopNamesLoD(LoD, 19000)
assert len(filtered) == 8
assert len(filtered[0]) == 4

assert len(LoD) == 10
assert len(LoD[0]) == 4
```

**Q6:** Write a function 

```
    addCatColumnLoD(tnLoD, threshold1, threshold2)
```

that adds a categorical column to a LoD representation in parameter `tnLoD` with the new column named `category` whose values are the strings `"small"` when `count` is below `threshold1`, is `"medium"` when `count` is greater than or equal to `threshold1` and less than `threshold2`, and `"large"` when `count` is greater than or equal to `threshold2`.

This change to tnLoD happens in place, rather than creating a new list, and so nothing is returned from the function.

<!-- grade_id: cell-0d9878d3fcbb29fa, points: 0, grade: false, solution: true -->
```python
def addCatColumnLoD(tnLoD, threshold1, threshold2):
    """ 
    This function adds a categorical column to tnLoD whose values are 
    "small" (count<threshold1), "medium" (threshold1<=count<threshold2), 
    "large" (count>=threshold2)
    """
    for data in range(len(tnLoD)):
        count = tnLoD[data]['count']
        tnLoD[data]['category'] = "small" if count < threshold1 else "medium" if count >= threshold1 and count < threshold2 else "large"
```

```python
# Debugging cell
LoD = [{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, 
       {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, 
       {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, 
       {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, 
       {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, 
       {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, 
       {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, 
       {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, 
       {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, 
       {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]

addCatColumnLoD(LoD, 19000, 19500)
print(LoD)
```

Output:
```
[{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936, 'category': 'large'}, {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305, 'category': 'medium'}, {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455, 'category': 'large'}, {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635, 'category': 'large'}, {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496, 'category': 'medium'}, {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117, 'category': 'medium'}, {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800, 'category': 'large'}, {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798, 'category': 'small'}, {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688, 'category': 'small'}, {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837, 'category': 'large'}]
```

<!-- grade_id: cell-1da7403c4d9cd8aa, points: 3, grade: true, solution: false -->
```python
# Testing cell
assert addCatColumnLoD.__doc__ is not None # don't forget the docstring!

LoD = [{'year': 2014, 'sex': 'Female', 'name': 'Emma', 'count': 20936}, 
       {'year': 2014, 'sex': 'Male', 'name': 'Noah', 'count': 19305}, 
       {'year': 2015, 'sex': 'Female', 'name': 'Emma', 'count': 20455}, 
       {'year': 2015, 'sex': 'Male', 'name': 'Noah', 'count': 19635}, 
       {'year': 2016, 'sex': 'Female', 'name': 'Emma', 'count': 19496}, 
       {'year': 2016, 'sex': 'Male', 'name': 'Noah', 'count': 19117}, 
       {'year': 2017, 'sex': 'Female', 'name': 'Emma', 'count': 19800}, 
       {'year': 2017, 'sex': 'Male', 'name': 'Liam', 'count': 18798}, 
       {'year': 2018, 'sex': 'Female', 'name': 'Emma', 'count': 18688}, 
       {'year': 2018, 'sex': 'Male', 'name': 'Liam', 'count': 19837}]

res = addCatColumnLoD(LoD, 19000, 19500)
assert res is None

assert len(LoD) == 10
assert isinstance(LoD[0], dict)
assert len(LoD[0]) == 5
assert 'category' in LoD[0]
assert [ rowD['category'] for rowD in LoD ] == ['large', 'medium', 'large',
 'large', 'medium', 'medium', 'large', 'small', 'small', 'large']
```

---

---
## Part C

How much time (in minutes/hours) did you spend on this homework assignment?

I spent around an hour and a half on this assignment.

Who was your partner for this assignment?  If you worked alone, say so instead.

I worked alone on this assignment.