# Denison CS181/DA210 Homework 1.e - Step 2

Before you turn this notebook in, make sure everything runs as expected. This is a combination of **restarting the kernel** and then **running all cells**.

Make sure you fill in any place that says `# YOUR CODE HERE` or "YOUR ANSWER HERE".

---

```python
import os
import os.path
```

---

## Part E

**Q11:** Use a list comprehension to implement a function 

```
    mult_by_2(data)
```

that multiplies each item in `data` by 2, returning the resulting list. For example, `mult_by_2([1,2,3])` returns `[2,4,6]`.

<!-- grade_id: cell-75d78706c6b457bf, points: 0, grade: false, solution: true -->
```python
def mult_by_2(data):
    """
    This function multiplies each item in data by 2
    """
    return [x*2 for x in data]
```

```python
# Debugging cell
mult_by_2([1,2,3]) # should be [2, 4, 6]
```

Output:
```
[2, 4, 6]
```

<!-- grade_id: cell-54db22cae09e3c07, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert mult_by_2.__doc__ is not None # make sure to have a docstring

assert mult_by_2([1]) == [2]
assert mult_by_2([1, 2]) == [2, 4]
assert mult_by_2([1,2,3,4,5,6,7,8,9,10]) == [2, 4, 6, 8, 10, 12, 14, 16, 18, 20]
assert mult_by_2([]) == []

mylist = [1, 2, -3]
assert mult_by_2(mylist) == [2, 4, -6]
assert mylist == [1, 2, -3]
```

**Q12:** Write a function

```
    add2odd(data)
```

that returns a new list based on `data` by including and adding 2 to every number that is odd, and leaving out any even numbers.

Use a list comprehension with a Boolean condition.  So, for instance, `add2odd([1,2,3,4,5])` results in the list `[3, 5, 7]`.

<!-- grade_id: cell-e45bf61a17020be8, points: 0, grade: false, solution: true -->
```python
def add2odd(data):
    """
    This function returns a new list based on data
    by including and adding 2 to every odd number
    and leaving out even numbers
    """
    return [x+2 for x in data if x %2 != 0]
```

```python
# Debugging cell
add2odd([1,2,3,4,5]) # should be [3, 5, 7]
```

Output:
```
[3, 5, 7]
```

<!-- grade_id: cell-f1afbb66fb298a2f, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert add2odd.__doc__ is not None # make sure to have a docstring

assert add2odd([]) == []
assert add2odd([1,2,3,4,5]) == [3, 5, 7]
assert add2odd([2,4,18,20,100,6]) == []
assert add2odd([1, 1, 1, 1, 1]) == [3, 3, 3, 3, 3]
```

**Q13:** Write a function

```
    get_tuples(data)
```

that uses a list comprehension to extract a list of tuples based on `data`, where the first item of each tuple is the index of the data, and the second item is the corresponding item in `data`.

For example, if `data = ["a", "b", "c"]` you return `[(0, 'a'), (1, 'b'), (2, 'c')]`.

Note that your solution **must** use a list comprehension.

<!-- grade_id: cell-b4370e23ea760754, points: 0, grade: false, solution: true -->
```python
def get_tuples(data):
    """
    This function uses a list comprehension to extract a
    list of tupls based on data where the first item
    of each tuple is the index of the data and the
    second item is the corresponding item in data
    """
    return [(data.index(item), item) for item in data]
```

```python
# Debugging cell
get_tuples(["a", "b", "c"]) # should be [(0, 'a'), (1, 'b'), (2, 'c')]
```

Output:
```
[(0, 'a'), (1, 'b'), (2, 'c')]
```

<!-- grade_id: cell-40e9d187d35a6b50, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert get_tuples.__doc__ is not None # make sure to have a docstring

assert get_tuples([]) == []
assert get_tuples(['a', 'b', 'c', 'd', 'e']) == [(0, 'a'), (1, 'b'), (2, 'c'), (3, 'd'), (4, 'e')]
assert get_tuples([0, 1, 2]) == [(0, 0), (1, 1), (2, 2)]
```

---

## Part F

**Q14:** Write a function

```
    filterOut(data, value)
```

that returns a new list that contains the same elements as the list `data` except for those that equal `value`. Use a list comprehension. Your function should filter out all items equal to `value`.

For example, `filterOut([3, 1, 5, 3, 9], 3)` should return the list `[1, 5, 9]`.

<!-- grade_id: cell-3f464b42be4a40e5, points: 0, grade: false, solution: true -->
```python
def filterOut(data, value):
    """
    This function returns a new list that contains
    the same elements as the list data except for those that
    equal value
    """
    return [x for x in data if x != value]
```

```python
# Debugging cell
filterOut([3, 1, 5, 3, 9], 3) # should be [1, 5, 9]
```

Output:
```
[1, 5, 9]
```

<!-- grade_id: cell-cfa882c037ff13ff, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert filterOut.__doc__ is not None # make sure to have a docstring

assert(filterOut([3,1,5,3,9], 1) == [3,5,3,9])
assert(filterOut([3,1,5,3,9], 9) == [3,1,5,3])
assert(filterOut([3,1,5,3,9], 3) == [1,5,9])
assert(filterOut([3,1,5,3,9], 2) == [3,1,5,3,9])
```

**Q15:** Write a function 

```
    filterOutPosition(data, index)
```

that returns a new list that contains the same elements as the list `data` except for the one at the given `index`. If the value of `index` is negative or exceeds the length of data, return a copy of the original list.  Use a list comprehension.

For example, `filterOutPosition([3, 1, 5, 9], 2)` should return the list `[3, 1, 9]`.

Hint: Your source list must contain all the needed information for the list comprehension to do its job.  Explore the `zip` function and/or the `enumerate` function.

<!-- grade_id: cell-e2a5560beae8497e, points: 0, grade: false, solution: true -->
```python
def filterOutPosition(data, index):
    """
    This function returns a new list that contains the same elements
    as the list data except for the one at the given index. If the value
    of index is negative or exceeds the length of data, it will return
    a copy of the original list
    """
    list_with_pos = list(enumerate(data))
    if (len(list_with_pos) < index + 1) or (index < 0):
        return data
    else:
        return [x for x_index, x in list_with_pos if x_index != index]
```

```python
# Debugging cell
filterOutPosition([3,1,5,9], 2) # should be [3, 1, 9]
```

Output:
```
[3, 1, 9]
```

<!-- grade_id: cell-7cd21bc67f04e583, points: 2, grade: true, solution: false -->
```python
# Testing cell
assert filterOutPosition.__doc__ is not None # make sure to have a docstring

assert(filterOutPosition([3, 1, 5, 9], 0) == [1, 5, 9])
assert(filterOutPosition([3, 1, 5, 9], 2) == [3, 1, 9])
assert(filterOutPosition([3, 1, 5, 9], 3) == [3, 1, 5])
assert(filterOutPosition([3, 1, 5, 9], -3) == [3, 1, 5, 9])
assert(filterOutPosition([3, 1, 5, 9], 4) == [3, 1, 5, 9])
```

---

---
## Part G

How much time (in minutes/hours) did you spend on this homework assignment?

I spent around 2 hours on this assignment.

Who was your partner for this assignment?  If you worked alone, say so instead.

I worked alone on this assignment.